---
title       : Themes and customizing graphs
description : Here you'll learn how to make your stuff nicer

---
## First exercise

```yaml
type: NormalExercise
key: e7844153e7
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
