---
title: 'Some exercises to make you strong'
description: 'Here are some exercises'
---

## Exercise 1

```yaml
type: NormalExercise
key: 5c1f8bba30
lang: r
xp: 100
skills: 1
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
