---
title: 'Some vizualization tips and best practices'
description: 'Insert the chapter description here'
---

## First Exercise

```yaml
type: NormalExercise
key: de07a5340a
lang: r
xp: 100
skills: 1
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
