---
title: Learning different geoms
description: >-
  Im quite happy the video works

---



## Stock Video Exercise

```yaml
type: VideoExercise 
lang: r
xp: 50 
video_link: //player.vimeo.com/video/160339466
video_hls: //videos.datacamp.com/transcoded/672_intermediate_r/v2/hls-ch1_1.master.m3u8
```







## Hall & Oates Video Exercise 

```yaml
type: VideoExercise
lang: r
xp: 50
video_link: "//player.vimeo.com/video/260397176"
key: 4dc7501f08
```



---
## Boxplots

```yaml
type: NormalExercise
key: c257f6f9e7
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Histogram

```yaml
type: NormalExercise
key: 5e80e714b5
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Density

```yaml
type: NormalExercise
key: dc9f0f3e66
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Smooth

```yaml
type: NormalExercise
key: d687abfe4f
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Jitter

```yaml
type: NormalExercise
key: 15a0ef0170
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Points again

```yaml
type: NormalExercise
key: d834e080e2
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Bars again

```yaml
type: NormalExercise
key: 2918999bc9
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
